# Bonjour ! 

```python
import sys
sys.path.append('../../../challenges_utils')
sys.path.append('../../../fl')
```

```python
from flags import flag
```

<!-- grade_id: cell-46dde148cd24d6f9, points: 0, grade: false, solution: true -->
```python
### BEGIN YOUR CODE HERE

### END YOUR CODE HERE
```

<!-- grade_id: cell-500dbee48bb2c260, points: 1, grade: true, solution: false -->
```python
assert bonsoir == 2
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[7], line 1
----> 1 assert bonsoir == 2

AssertionError:
```

```python

```